# CS 3110/5990: Data Privacy
## Homework 4

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

## Question 1 (10 points)

Complete the definition of `dp_sum_capgain` below. Your definition should compute a differentially private sum of the "Capital Gain" column of the `adult` dataset, and have a total privacy cost of `epsilon`.

<!-- grade_id: cell-6a369f61d8eef943, points: 0, grade: false, solution: true -->
```python
def dp_sum_capgain(epsilon):
    clip = 2**17
    clipped = adult['Capital Gain'].clip(lower=0, upper=clip)
    return laplace_mech(clipped.sum(), clip, epsilon)

dp_sum_capgain(1.0)
```

Output:
```
34803235.1551215
```

<!-- grade_id: cell-e5b1229ba2249f71, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 1

real_sum = adult['Capital Gain'].sum()
r1 = np.mean([pct_error(real_sum, dp_sum_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_sum, dp_sum_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_sum, dp_sum_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 10
assert r2 < 2
assert r3 < 0.2
```

Output:
```
Average errors: 3.2366065175259178 0.35134504168713204 0.03889186947158548
```

## Question 2 (10 points)

In 2-5 sentences each, answer the following:

- What clipping parameter did you use in your definition of `dp_sum_capital`, and why?
- What was the sensitivity of the query you used in `dp_sum_capital`, and how is it bounded?
- Argue that your definition of `dp_sum_capital` has a total privacy cost of `epsilon`

<!-- grade_id: cell-986eec754f32c072, points: 10, grade: true, solution: true -->
My upper clipping parameter was set to 2<sup>17</sup>, which I arrived at through trial and error. The sensitivity of the query is 2<sup>17</sup> because the lower bound is 0 and the sensitivity of a clipped query is the difference between the upper and lower bounds. The total privacy cost for the function is equal to epsilon because ony one query with a total privacy cost of epsilon is being made

## Question 3 (10 points)

Complete the definition of `dp_avg_capgain` below. Your definition should compute a differentially private average (mean) of the "Capital Gain" column of the adult dataset, and have a **total privacy cost of epsilon**.

<!-- grade_id: cell-93f2a5153dca0e95, points: 0, grade: false, solution: true -->
```python
def dp_avg_capgain(epsilon):
    sum_capgain = dp_sum_capgain(epsilon/2)
    count_capgain = laplace_mech(adult.shape[0], 1, epsilon/2)
    return sum_capgain / count_capgain

dp_avg_capgain(1.0)
```

Output:
```
1103.5480936663057
```

<!-- grade_id: cell-ef77b325b5c58908, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3

real_avg = adult['Capital Gain'].mean()
r1 = np.mean([pct_error(real_avg, dp_avg_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_avg, dp_avg_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_avg, dp_avg_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 20
assert r2 < 4
assert r3 < 0.4
```

Output:
```
Average errors: 7.177193846321342 0.8218888222881664 0.06780609292405845
```

## Question 4 (10 points)

In 2-5 sentences each, answer the following:

- Argue that your definition of `dp_avg_capgain` has a total privacy cost of `epsilon`
- For sums and averages, which seems to be more important for accuracy - the value of the clipping parameter $b$ or the scale of the noise added? Why?
- Do you think the answer to the previous point will be true for every dataset? Why or why not?

<!-- grade_id: cell-7ea16016366529bd, points: 10, grade: true, solution: true -->
The total privacy cost of `dp_avg_capgain` is `epsilon` because two queries equal to epsilon/2 are made. The total privacy cost can be calculated using sequential composition, bringing us to a total privacy cost of epsilon. 

## Question 5 (20 points)

Write a function `auto_avg` that returns the differentially private average of a Pandas series `s`. Your function should automatically determine the clipping parameter `b`, and should enforce differential privacy for a **total privacy cost** of `epsilon`. You can assume that all values are non-negative (i.e. 0 or greater).

<!-- grade_id: cell-10b352b38a9b5f85, points: 0, grade: false, solution: true -->
```python
def auto_avg(s, epsilon):
    clips = [10**n for n in range(10)]
    n = len(clips)
    
    # Internally redefine epsilon to satisfy DP
    epsilon = epsilon / (n+1)
    
    noisy_count = laplace_mech(len(s), 1, epsilon)
    last = -1
    for clip in clips:
        noisy_sum = laplace_mech(s.clip(lower=0, upper=clip).sum(), clip, epsilon)
        noisy_avg = noisy_sum / noisy_count
        
        # Difference of less than a 10th of the clip -- 100th didn't seem to work
        if abs(noisy_avg - last) < .01 * clip:
            return last
        else:
            last = noisy_avg

```

<!-- grade_id: cell-49fb57b1ebfcc0da, points: 20, grade: true, solution: false -->
```python
# TEST CASE for question 5
r1 = np.mean([pct_error(adult['Age'].mean(), auto_avg(adult['Age'], 1.0)) for _ in range(20)])
r2 = np.mean([pct_error(adult['Capital Gain'].mean(), auto_avg(adult['Capital Gain'], 1.0)) for _ in range(20)])
r3 = np.mean([pct_error(adult['fnlwgt'].mean(), auto_avg(adult['fnlwgt'], 1.0)) for _ in range(20)])

print('Average errors:', r1, r2, r3)
assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 1
assert r2 < 100
assert r3 < 1
```

Output:
```
Average errors: 0.07470988494570033 51.01695671480552 0.14172860161702638
```

## Question 6

In 2-5 sentences each, answer the following:

- Explain your strategy for implementing `auto_avg`
- Argue informally that your definition of `auto_avg` has a total privacy cost of `epsilon`
- Did your solution work well for all three example columns? If it did not work well on any of them, why not?
- When is your solution likely to *not* work well? (i.e. what properties does the data have to have, in order for your solution to not work well?)

<!-- grade_id: cell-0ff63739ed66adf7, points: 10, grade: true, solution: true -->
- My solution creates a set of clipping parameters and iterates over them until it finds a clip that creates relatively little change from the last dataset. When there's relatively little change the result from the previous clip is accepted.
- My definition of `auto_avg` has a total privacy cost of `epsilon` because the algorithm makes at most *n* iterations at `epsilon / n` privacy.
- My solution did not work well for Capital Gain. Looking at the assertions it doesn't seem like it was meant to, but I'm not sure why. Maybe it clipped too tight?
- The solution has a fixed set of clipping parameters - any data value larger than the greatest clip will be ignored. The solution is also vulnerable to large gaps in the dataset.