# CS 3110/5110: Data Privacy
## Homework 1

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
adult_pii = adult[['Name', 'DOB', 'SSN', 'Zip', 'Age']]
adult_deid = adult.drop(columns=['Name', 'SSN'])
```

## Question 1 (20 points)

Using the dataframes `adult_pii` and `adult_deid`, write code to conduct a linking attack to recover the names of as many individuals in `adult_deid` as possible. Your solution should be parameterized by the set of columns to use in the attack.

<!-- grade_id: cell-e6389ceace711c32, points: 0, grade: false, solution: true -->
```python
def linking_attack(cols):
    return pd.merge(adult_pii, adult_deid, left_on=cols, right_on=cols)
```

<!-- grade_id: cell-251b9f1c6f7b4630, points: 20, grade: true, solution: false -->
```python
# TEST CASES for Question 1

assert len(linking_attack(['Zip'])) == 43191
assert len(linking_attack(['Zip', 'DOB'])) == 32563
assert len(linking_attack(['Zip', 'Age'])) == 32755
```

## Question 2 (5 points)

How many individuals in this dataset are uniquely identified by their Zip code? How many are uniquely identified by their age?

Hint: note that the number of *unique ZIP codes* is **different** from the number of *individuals uniquely identified by ZIP code*.

Hint: you can use the `value_counts` method to count the number of occurences of each value in a series.

<!-- grade_id: cell-0e380ac68dd0c77e, points: 0, grade: false, solution: true -->
```python
def unique_zipcode():
    # YOUR CODE HERE
    raise NotImplementedError()

def unique_dob():
    # YOUR CODE HERE
    raise NotImplementedError()
```

Output:
```
File ~\anaconda3\envs\CS3110-Data-Privacy\Lib\site-packages\pandas\core\series.py:1116, in Series._get_value(self, label, takeable)
   1113     return self._values[label]
   1115 # Similar to Index.get_value, but we do not fall back to positional
-> 1116 loc = self.index.get_loc(label)
   1118 if is_integer(loc):
   1119     return self._values[loc]

File ~\anaconda3\envs\CS3110-Data-Privacy\Lib\site-packages\pandas\core\indexes\base.py:3655, in Index.get_loc(self, key)
   3653     return self._engine.get_loc(casted_key)
   3654 except KeyError as err:
-> 3655     raise KeyError(key) from err
   3656 except TypeError:
   3657     # If we have a listlike key, _check_indexing_error will raise
   3658     #  InvalidIndexError. Otherwise we fall through and re-raise
   3659     #  the TypeError.
   3660     self._check_indexing_error(key)

KeyError: 1
```

<!-- grade_id: cell-9b009427c61dbb2b, points: 5, grade: true, solution: false -->
```python
# TEST CASES (INCOMPLETE) for Question 2
# INCOMPLETE means that additional test cases will be used to check your answer during grading

assert unique_zipcode() > 0
assert unique_dob() > 0

```

Output:
```
Cell In[13], line 4
      1 # TEST CASES (INCOMPLETE) for Question 2
      2 # INCOMPLETE means that additional test cases will be used to check your answer during grading
----> 4 assert unique_zipcode() > 0
      5 assert unique_dob() > 0

Cell In[12], line 3, in unique_zipcode()
      1 def unique_zipcode():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

## Question 3 (10 points)

Write code to determine the `Education-Num` of the individual named Ardyce Golby by performing a differencing attack. Your code should *only* use aggregate data to find Ardyce's education number.

<!-- grade_id: cell-ef28e136429ae8c2, points: 0, grade: false, solution: true -->
```python
def ardyce_education():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d06fb722893a39a5, points: 10, grade: true, solution: false -->
```python
# TEST CASE for Question 3
assert ardyce_education() == 12
```